```python
%%capture
# Install the autograder to check your answers as you go along
import sys
!{sys.executable} -m pip install otter-grader
# Initialize Otter 
import otter
grader = otter.Notebook("group_activity5.ipynb")
```

# In-class coding exercise #5
Objective: Practice using Pandas and the basics of Xarray

## Introduction to group coding exercises
Today you’ll work on this exercise in the same groups of 3-4 you were assigned last week, submitting a single notebook file at the end of the class period. Decide amongst yourselves which member will upload the completed notebook to Gradescope this week. Make sure that everyone takes a turn being the “Uploader”. _You cannot upload the final code two weeks in a row._

Designate a different group member to be the "Reporter". The Reporter will be in charge of participating in the group discussion at the end of the class session.

### Workflow
Each question will be timed to ensure that everyone gets to work on at least a part of every question. Group activities are not graded by completeness or correctness, but by effort. We will be breaking down each question in the following order:  
1. Independent work 
2. Group work and discussion on coding question
3. Group work and discussion on reflection questions

You are welcome and encouraged to communicate with other groups and the teaching team when you feel stuck on a problem. 

As a reminder, we will be grading based best practices in coding. These include: 
1) Variables are used to store objects

2) Code is commented adequately

3) Variables are names appropriately

4) Code is efficient with minimal unnecessary lines 

5) Documenting help from outside sources, such as from other groups or online documentation. 

6) Final notebook fully runs from start to finish. A good way to check this is restarting the kernel and fully running through all the cells to check for any errors.

### Storing your answers
In the code cells where you will write your answers, there will be comments denoting:

"**# your code**"

and 

"**# answer variables**"

You may store any intermediary variables in the **your code** section. If you do not have any intermediary, you can also store your answer directly in the answer variables.

### Required Plot Elements for Figures
This assignments requires you to create and design figures using `matplotlib`. To practice good plotting practices, each figure will require the following to receive full points:
1) Concise, descriptive title for each figure/subplot
2) Axis labels with units (when possible)
3) Appropriate axis limits (minimum and maximum)
4) Appropriate tick resolution
5) Legend when using different datasets 
6) Appropriate font size (a good range is 12-15)

## Note here **and in the Gradescope submission** each of your group members:
1.
2. 
3.

# Q1: Practice with `pandas`: Ballard Locks salmon counts (50 minutes)

## Part 1:
### <font color='red'><i>~10-15 minutes</i></font>
Here are some local salmon data we can load and visualize using pandas! Let's first get an overview of the data by practicing some dataframe functions.

We will load this data from the URL listed below in the starter code cell, stored in the `filepath` string. Instead of downloading the data itself into the Jupyter environment, we will practice pulling data into our environment without saving it to the server (disk). 

### Instructions
1. Load the salmon data URL into Pandas using the pandas `read_csv()` function and store it into the `salmon_df` variable. Follow the steps below to set some positional arguments in the read_csv() function while loading the file:
> * Specify that the first (0th) column (the dates) should be the index.
>
> * Also specify that Pandas should parse the index as dates (datetimes).
>
> * Consult the documentation for [`pd.read_csv()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.read_csv.html) to find the arguments to specify these two things.
2. Display the data.
3. Use `.describe()` to view the summary statistics.
4. Answer the following questions with your group in a comment or a markdown cell below your code:

* How many salmon species are counted?
* When does this data start and end?
* What are the mean daily counts for each species?
* What are the maximum daily counts for each species?

```python
# starter code cell, DO NOT CHANGE THIS CODE!
filepath = 'https://raw.githubusercontent.com/OCEAN-215-CLASS/week6_inclass_data/main/ballard_salmon_counts.csv'
```

_Type your answer here, replacing this text._

```python
# what packages do you need to import?

# your code

# print/display statements
```

```python
grader.check("Question 1, Part 1")
```

<!-- BEGIN QUESTION -->

## Part 2: **Plot the data!** 
### <font color='red'><i>~15-20 minutes</i></font>

### Instructions
1. Create a single blank figure using `matplotlib.pyplot subplots()` or `figure()`.  Set the `figsize` argument so it will take up the entire width of the page. 
2. Use either `plt.plot()` or `ax.plot()` (depending on how you created your blank figure in Part 1) to make line plots of each of the three species' counts over time. In other words, the x-values should be datetimes from the index and the y-values should be daily salmon counts for each species. Your plot should have 3 different lines.
3. Make each species the following colors for each line on your plot:
> * Chinook: 'forestgreen'
>
> * Coho: 'darkcyan'
>
> * Sockeye: 'salmon'
>


4. Label your plot axes and add a title.
5. Add a grid to your plot.
6. Add a legend.
7. Double check that your plot contains all the required elements as listed in the Assignment Instructions above.

```python
# create your figure 
...
```

<!-- END QUESTION -->

### Part 3: **Index into the data to discover information** (2 points)
### <font color='red'><i>~15-20 minutes</i></font>

1. Use indexing to find out how many coho salmon passed through Ballard Locks on September 30, 2020, the day that Autumn Quarter started at UW in year 2020. Store this number in the `coho_count` variable without hardcoding. 
> _Hint:_ you'll need `.loc[]` and a conditional to do this
2. Use indexing with slicing (`:`) to get the sockeye salmon counts for all dates in the years specified below. Make sure to print the results to check that they're correct.
> * First, do this for 2020. Save the result as a new variable, `sockeye_2020`.
>
> * Then, do the same for 2013. Save this as `sockeye_2013`.
>

3. Apply `NumPy` functions to `sockeye_2020` and `sockeye_2013` to find the following and store in their corresponding answer variables:
> * The maximum daily sockeye count in each year (2013 and 2020).
>
> * The total number of sockeye that passed through Ballard Locks in each year.
>


**OPTIONAL (only complete if you have extra time)**

4. Do Steps 2-3 again, with a single line of code for each year and `numpy` function— in other words, without saving the sliced Pandas objects as intermediary variables. This will result in 1 line of code to slice and apply the function for each year and function.
> * Print your outputs. They should match your answer variables in Part 3 if correct.

```python
# Solution cell for Part 3
# which packages do you need to import? 

# your code

# answer variables
coho_count = ""
sockeye_2020 = ""
sockeye_2013 = ""
max_2020 = ""
max_2013 = ""
total_2020 = ""
total_2013 = ""

# code for optional code if you have time

```

```python
grader.check("Question 1, Part 3")
```

<!-- BEGIN QUESTION -->

# Reflection questions (5 minutes)
The purpose of the reflection is to inform us as instructors about students comfort level with course content. We use these answers to inform how we spend class time and design coursework in subsequent weeks. This question is graded for completeness, so please answer each question in the text box below. Be concise in your answers (max. 2 sentences). 

1) What do you feel you excelled at in this exercise? Why?

2) What did you struggle with most in the exercise? Why?

3) Is there any section of the question that you did not complete? If so, briefly describe why and the section you spent the most time on. 

4) Is there any topic you feel we need to revisit or review in class? Why?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

# Question 2: Practice reading files with `xarray`: World Ocean Atlas global ocean temperatures
### <font color='red'><i>Max time for all Q2 parts: 30 minutes</i></font>

1. Load the WOA18 netCDF URL into xarray using `xr.open_dataset()` using the url string in the cell below. Set the named argument for "decode_times=False". This will keep the dates as a string instead of converting them to datetimes. 
2. Display the data.
3. Using just the interactive display, answer the following questions with your group:

> 1. How many data variables are there?
> 2. The variable abbreviations aren't very informative. Using the page icon, find out what at least 3 variables represent and list them.
> 3. What is the time range of the data?
> 4. What is the latitude and longitude resolution (spacing) of the data? Note that we call this the "grid spacing" or "resolution" of the data.

**Optional if you have time**: <br>
Subset your xarray dataset values using xarray or numpy functions to answer the following extra credit questions. Type your answers in the markdown cell below Part 3. Answer the questions in each part in comments below or above your code in the solution cell.
> 5. What are the dimensions of the dataset? Type this in a new variable to store and print it. Take a look at your coordinates, specifically depth. Use a `numpy` function to calculate the deepest measurement of depth. 
> 6. Let's take a look at how much sampling was done at the equator. Subset your dataset's "lat" coordinates to 0. Do you see any measurements? 
- If not, try experimenting with the "method" argument to see if you can see other measurements nearby. How many measurements do you see now?
- What is the range of values for lat (latitude) you see shown when you use "method"?

> 7. Create a plot of temperature at the equator using the `plot()` function on your subsetted equatorial dataset. Don't forget to reverse the y-axis for depth!
> 8. How does temperature vary at different latitudes? Now, try subsetting a dataset to -60 with the "lat" coordinates. Plot your figure and compare it against your equatorial plot.

```python
# Starter code cell, DO NOT CHANGE!
url="https://www.ncei.noaa.gov/thredds-ocean/dodsC/ncei/woa/temperature/decav/1.00/woa18_decav_t00_01.nc"
```

_Type your answer here, replacing this text._

```python
# what packages do you need to import

# your code

# display your dataframe

```

```python
## OPTIONAL SECTION 
# your code

# your plots 

...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

# Reflection questions (5 minutes)
The purpose of the reflection is to inform us as instructors about students comfort level with course content. We use these answers to inform how we spend class time and design coursework in subsequent weeks. This question is graded for completeness, so please answer each question in the text box below. Be concise in your answers (max. 2 sentences). 

1) What do you feel you excelled at in this exercise? Why?

2) What did you struggle with most in the exercise? Why?

3) Is there any section of the question that you did not complete? If so, briefly describe why and the section you spent the most time on. 

4) Is there any topic you feel we need to revisit or review in class? Why?

_Type your answer here, replacing this text._

<!-- END QUESTION -->



Tests for Question 1, Part 1 (points: 0)
```
>>> # test if columns and indices are correct
>>> (list(salmon_df.columns) == ['Chinook', 'Coho', 'Sockeye']) and (salmon_df.index.dtype == '<M8[ns]')
True
```
failure_message: Check your salmon_df columns or indices!
```
>>> # test if shape of salmon_df is correct (are all the data there?)
>>> salmon_df.shape == (419, 3)
True
```
failure_message: Check your salmon_df shape!


Tests for Question 1, Part 3 (points: 0)
```
>>> # test if columns and indices are correct
>>> len(sockeye_2020.loc[sockeye_2020.notnull()]) == 64
True
```
failure_message: Check your sockeye_2020 variable!
```
>>> # test if columns and indices are correct
>>> len(sockeye_2013.loc[sockeye_2013.notnull()]) == 56
True
```
failure_message: Check your sockeye_2013 variable!
```
>>> # test if columns and indices are correct
>>> int(max_2020) == 1961
True
```
failure_message: Check your sockeye_2020 variable!
```
>>> # test if columns and indices are correct
>>> int(max_2013) == 12936
True
```
failure_message: Check your max_2013 variable!
```
>>> # test if columns and indices are correct
>>> int(total_2020) == 22954
True
```
failure_message: Check your sum_2020 variable!
```
>>> # test if columns and indices are correct
>>> int(total_2013) == 178422
True
```
failure_message: Check your sum_2013 variable!
